# Homework 5

In this homework your goal is to analyze the shooting percentage of a given player over the course of the season. 


## Problem 1

In this problem, you are asked to implement a function to get the player's shots data. In particular, your function should read in "data/shot_logs.csv" and filter the dataset to return the player's data. 




<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
import pandas as pd


def GetPlayerData(player):
    df_shots_player = pd.read_csv("data/shot_logs.csv", cloums=)
    df_shots_player=df_shots_player[df_shots_player['player_name']==player]    
    return df_shots_player

```

<!-- grade_id: cell-2853176972493b6e, points: 25, grade: true, solution: false -->
```python
df_shots_player = GetPlayerData("stephen curry")
assert df_shots_player.shape == (968, 21)
df_shots_player = GetPlayerData("brian roberts")
assert df_shots_player.shape == (372, 21)
```

## Problem 2

You need to implement a function to add a Month column to the existing data, which contains the month from the MATCHUP column. 

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def AddMonth(df_shots_player):
    d = {'JAN':1, 'FEB':2, 'MAR':3, 'APR':4, 'MAY':5, 'JUN': 6, 'JUL':7,'AUG':8,'SEP':9,'OCT':10,'NOV':11,'DEC':12 }
    Month=[]
    for index,  row in df_shots_player.iterrows():
        time=row['MATCHUP']
        for m in d.keys():
            if str(m)==str(time[:3]):
                Month.append(d[m])
    df_shots_player['Month']=Month
    return df_shots_player
AddMonth(df_shots_player)
```

Output:
```
GAME_ID                     MATCHUP LOCATION  W  FINAL_MARGIN  \
0    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
1    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
2    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
3    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
4    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
5    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
6    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
7    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
8    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
9    21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
10   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
11   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
12   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
13   2140088
```

```python
# alternative method
def AddMonth(df_shots_player):
    d = {'JAN':1, 'FEB':2, 'MAR':3, 'APR':4, 'MAY':5, 'JUN': 6, 'JUL':7,'AUG':8,'SEP':9,'OCT':10,'NOV':11,'DEC':12 }
    def monthFunc(item):
        month=item.split(' ')[0]
        return d[month]
    df_shots_player['Month']=df_shots_player.MATCHUP.apply(monthFunc)
    return df_shots_player
AddMonth(df_shots_player)
```

Output:
```
GAME_ID                     MATCHUP LOCATION  W  FINAL_MARGIN  \
0    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
1    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
2    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
3    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
4    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
5    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
6    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
7    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
8    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
9    21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
10   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
11   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
12   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
13   2140088
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 25, grade: true, solution: false -->
```python
df_shots_player = GetPlayerData("stephen curry")
df_shots_player = AddMonth(df_shots_player)
assert round(df_shots_player["Month"].describe()["mean"], 2) == 6.65

df_shots_player = AddMonth(GetPlayerData("brian roberts"))
assert round(df_shots_player["Month"].describe()["mean"], 2) == 6.69
```

## Problem 3

You need to implement a function to add a column called "Short_Type" to df_shots_player that identifies the shot type as one of the three that we are interested in: "short" (<=7 ft.), "mid-range" (Any 2-pointer longer than 7 ft.) and "3-pointer". I would recommend using apply with a custom function. You need to rely on the column SHOT_DIST and PTS_TYPE.


<!-- grade_id: cell-a3c8e33b23a52cfd, points: 0, grade: false, solution: true -->
```python
def GetShotType(row):
    """
    This function takes a row of the dataset and returns
    "3-pointer", "short" or "mid-range" depending on the shot
    type.
    """
    
    if row['SHOT_DIST']<=7 and row['PTS_TYPE'] != 3: # some shot distance shorter than 7 but got 3 point
         return_type = 'short'
    elif row['SHOT_DIST']>7 and row['PTS_TYPE']==2:
         return_type = 'mid-range'
    else:
        return_type = "3-pointer"
    return return_type

def AddShotType(df_shots_player):
    df_shots_player['Shot_Type']=df_shots_player.apply(GetShotType,axis=1)
    return df_shots_player

AddShotType(df_shots_player)

```

Output:
```
GAME_ID                     MATCHUP LOCATION  W  FINAL_MARGIN  \
0    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
1    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
2    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
3    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
4    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
5    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
6    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
7    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
8    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
9    21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
10   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
11   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
12   21400890  MAR 03, 2015 - CHA vs. LAL        H  W             1   
13   2140088
```

<!-- grade_id: cell-d143e38c322d9286, points: 25, grade: true, solution: false -->
```python
import pandas as pd

df_shots_player = GetPlayerData("stephen curry")
df_shots_player = AddMonth(df_shots_player)
df_shots_player = AddShotType(df_shots_player)

assert df_shots_player[df_shots_player["Shot_Type"] == "3-pointer"].shape[0] == 456
assert df_shots_player[df_shots_player["Shot_Type"] == "short"].shape[0] == 224
assert df_shots_player[df_shots_player["Shot_Type"] == "mid-range"].shape[0] == 288
```

## Problem 4
You need to implement a function to compute the player's field goal percentage for each month and for each type of shot. Only look at months where he took at least 30 shots in each category. The function should return a dataframe for "stephen curry" as the follows:

<img src="Data/dataframe.png" width=400>

<!-- grade_id: cell-72cf7018f4d695a2, points: 0, grade: false, solution: true -->
```python
def ComputeGoalType(df_shots_player):
    d={}
    for i in range(1,13):
        mondf=df_shots_player.loc[df_shots_player['Month']==i] # get each month's data
        
        p3df=mondf.loc[mondf['Shot_Type']=='3-pointer']
        p3shotnum=p3df.Month.count() # the total shot number of 3 pointer
        
        midrangedf=mondf.loc[mondf['Shot_Type']=='mid-range']
        midrangeshotnum=midrangedf.Month.count()# the total shot number of mid-range
        
        shortdf=mondf.loc[mondf['Shot_Type']=='short']
        shortshotnum=shortdf.Month.count()# the total shot number of short
        
        if p3shotnum>=30 and midrangeshotnum>=30 and shortshotnum>=30:
            p3made=p3df.loc[p3df['SHOT_RESULT']=='made']
            p3madenum=p3made.Month.count()
            p3_perc_by_month=p3madenum/p3shotnum
            
            midrangemade=midrangedf.loc[midrangedf['SHOT_RESULT']=='made']
            midrangemadenum=midrangemade.Month.count()
            midrange_perc_by_month=midrangemadenum/midrangeshotnum
            
            shortmade=shortdf.loc[shortdf['SHOT_RESULT']=='made']
            shortmadenum=shortmade.Month.count()
            short_perc_by_month=shortmadenum/shortshotnum
            d[i]=[p3_perc_by_month,midrange_perc_by_month,short_perc_by_month]

  
    FG_perc_by_month=pd.DataFrame.from_dict(d,orient="index")
    FG_perc_by_month.columns = ["3-pointer","mid-range","short"]
    FG_perc_by_month.index.name="Month"
    
    return FG_perc_by_month
ComputeGoalType(df_shots_player)
```

Output:
```
3-pointer  mid-range     short
Month                                
1       0.408696   0.453333  0.650000
2       0.469136   0.346939  0.648649
11      0.433628   0.506329  0.647059
12      0.352381   0.478873  0.704918
```

<!-- grade_id: cell-06c7dc215450e8bd, points: 25, grade: true, solution: false -->
```python
df_shots_player = GetPlayerData("stephen curry")
df_shots_player = AddMonth(df_shots_player)
df_shots_player = AddShotType(df_shots_player)
FG_perc_by_month = ComputeGoalType(df_shots_player)
assert [round(x, 2) for x in FG_perc_by_month["3-pointer"].tolist()] == [0.41, 0.47, 0.43, 0.35]
assert [round(x, 2) for x in FG_perc_by_month["short"].tolist()] == [0.65, 0.65, 0.65, 0.70]
```