# Homework 5

In this homework your goal is to analyze the shooting percentage of a given player over the course of the season. 


## Problem 1

In this problem, you are asked to implement a function to get the player's shots data. In particular, your function should read in "data/shot_logs.csv" and filter the dataset to return the player's data. 




<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
import pandas as pd
  

def GetPlayerData(player):
    shots_player = pd.read_csv("data/shot_logs.csv", sep=",", header=0)
    filters = shots_player.get('player_name') == player
    
    result_player = shots_player[filters]    
    return result_player
# 
df_shots_player = GetPlayerData("stephen curry")
df_shots_player.shape
```

Output:
```
(968, 21)
```

<!-- grade_id: cell-2853176972493b6e, points: 25, grade: true, solution: false -->
```python
df_shots_player = GetPlayerData("stephen curry")
assert df_shots_player.shape == (968, 21)
df_shots_player = GetPlayerData("brian roberts")
assert df_shots_player.shape == (372, 21)
```

## Problem 2

You need to implement a function to add a Month column to the existing data, which contains the month from the MATCHUP column. 

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def AddMonth(df_shots_player):
    d = {'JAN':1, 'FEB':2, 'MAR':3, 'APR':4, 'MAY':5, 'JUN': 6, 'JUL':7,'AUG':8,'SEP':9,'OCT':10,'NOV':11,'DEC':12 }
    month = []
    for _, row in df_shots_player.iterrows():
        matchup = row.get('MATCHUP')  # "MAR 04, 2015 - CHA @ BKN"
        for key in d.keys():
            if key == matchup.split(' ')[0]:
                value = d.get(key)
                if value:
                    month.append(value)
    df_shots_player['Month'] = month
    return df_shots_player
AddMonth(df_shots_player)
```

Output:
```
GAME_ID                     MATCHUP LOCATION  W  FINAL_MARGIN  \
0    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
1    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
2    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
3    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
4    21400899    MAR 04, 2015 - CHA @ BKN        A  W            24   
..        ...                         ...      ... ..           ...   
367  21400028  NOV 01, 2014 - CHA vs. MEM        H  L            -2   
368  21400004  OCT 29, 2014 - CHA vs. MIL        H  W             2   
369  21400004  OCT 29, 2014 - CHA vs. MIL        H  W             2   
370  21400004  OCT 29, 2014 - CHA vs. MIL        H  W             2   
371  21400004  OCT 29, 2014 - CHA vs. MIL        H  W             2   

     SHOT_NUMBER  PERIOD GAME_CLOCK  SHOT_CLOCK  DRIBBLES  ...  PTS_TYPE  \
0              1       1       1:09        10.8         2  ...         2   
1
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 25, grade: true, solution: false -->
```python
df_shots_player = GetPlayerData("stephen curry")
df_shots_player = AddMonth(df_shots_player)
assert round(df_shots_player["Month"].describe()["mean"], 2) == 6.65

df_shots_player = AddMonth(GetPlayerData("brian roberts"))
assert round(df_shots_player["Month"].describe()["mean"], 2) == 6.69
```

## Problem 3

You need to implement a function to add a column called "Short_Type" to df_shots_player that identifies the shot type as one of the three that we are interested in: "short" (<=7 ft.), "mid-range" (Any 2-pointer longer than 7 ft.) and "3-pointer". I would recommend using apply with a custom function. You need to rely on the column SHOT_DIST and PTS_TYPE.


<!-- grade_id: cell-a3c8e33b23a52cfd, points: 0, grade: false, solution: true -->
```python
def GetShotType(row):
    """
    This function takes a row of the dataset and returns
    "3-pointer", "short" or "mid-range" depending on the shot
    type.
    """
    return_type = "3-pointer"
    
    if row.get('SHOT_DIST') > 7:
        if row.get('PTS_TYPE') == 2:
            return_type = 'mid-range'
    elif row.get('SHOT_DIST') <= 7:
        if row.get('PTS_TYPE') != 3: # some shot distance shorter than 7 but got 3 point
             return_type = 'short'
    return return_type

def AddShotType(df_shots_player):
    series_shot_type = df_shots_player.apply(GetShotType, axis=1)
    df_shots_player['Shot_Type'] = series_shot_type
    return df_shots_player

AddShotType(df_shots_player)
```

Output:
```
GAME_ID                     MATCHUP LOCATION  W  FINAL_MARGIN  \
14054  21400907  MAR 04, 2015 - GSW vs. MIL        H  W             9   
14055  21400907  MAR 04, 2015 - GSW vs. MIL        H  W             9   
14056  21400907  MAR 04, 2015 - GSW vs. MIL        H  W             9   
14057  21400907  MAR 04, 2015 - GSW vs. MIL        H  W             9   
14058  21400907  MAR 04, 2015 - GSW vs. MIL        H  W             9   
...         ...                         ...      ... ..           ...   
15017  21400014    OCT 29, 2014 - GSW @ SAC        A  W            18   
15018  21400014    OCT 29, 2014 - GSW @ SAC        A  W            18   
15019  21400014    OCT 29, 2014 - GSW @ SAC        A  W            18   
15020  21400014    OCT 29, 2014 - GSW @ SAC        A  W            18   
15021  21400014    OCT 29, 2014 - GSW @ SAC        A  W            18   

       SHOT_NUMBER  PERIOD GAME_CLOCK  SHOT_CLOCK  DRIBBLES  ...  SHOT_RESULT  \
14054            1       1       9:49        16.3 
```

<!-- grade_id: cell-d143e38c322d9286, points: 25, grade: true, solution: false -->
```python
import pandas as pd

df_shots_player = GetPlayerData("stephen curry")
df_shots_player = AddMonth(df_shots_player)
df_shots_player = AddShotType(df_shots_player)

assert df_shots_player[df_shots_player["Shot_Type"] == "3-pointer"].shape[0] == 456
assert df_shots_player[df_shots_player["Shot_Type"] == "short"].shape[0] == 224
assert df_shots_player[df_shots_player["Shot_Type"] == "mid-range"].shape[0] == 288
```

## Problem 4
You need to implement a function to compute the player's field goal percentage for each month and for each type of shot. Only look at months where he took at least 30 shots in each category. The function should return a dataframe for "stephen curry" as the follows:

<img src="Data/dataframe.png" width=400>

<!-- grade_id: cell-72cf7018f4d695a2, points: 0, grade: false, solution: true -->
```python
def ComputeGoalType(df_shots_player):
    FG_perc_by_month = None
    df_by_month = dict()
    for month in [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12]:
        # get each month's data
        filters = df_shots_player.get('Month') == month
        df_month = df_shots_player.loc[filters] 
        
        # the total shot number of 3 pointer
        filters = df_month.get('Shot_Type') == '3-pointer'
        df_3p = df_month.loc[filters]
        shotnum_3p = df_3p.Month.count() 
        
        # the total shot number of mid-range
        filters = df_month.get('Shot_Type') == 'mid-range'
        df_mid_range = df_month.loc[filters]
        shotnum_mid_range=df_mid_range.Month.count()
        
        # the total shot number of short
        filters = df_month.get('Shot_Type') == 'short'
        df_short = df_month.loc[filters]
        shotnum_short = df_short.Month.count()
        
        critical_value = 30
        if shotnum_3p >= critical_value:
            if shotnum_mid_range >= critical_value:
                if shotnum_short >= critical_value:
                    filters = df_3p.get('SHOT_RESULT') == 'made'
                    made_3p = df_3p.loc[filters]
                    madenum_3p = made_3p.Month.count()
                    p3_perc_by_month = madenum_3p / shotnum_3p

                    # 
                    filters = df_mid_range.get('SHOT_RESULT') == 'made'
                    mid_range_made = df_mid_range.loc[filters]
                    mid_range_madenum = mid_range_made.Month.count()
                    midrange_perc_by_month = mid_range_madenum / shotnum_mid_range

                    # 
                    filters = df_short.get('SHOT_RESULT') == 'made'
                    short_made = df_short.loc[filters]
                    short_made_num = short_made.Month.count()
                    short_perc_by_month = short_made_num / shotnum_short
                    
                    df_by_month[month] = [p3_perc_by_month, midrange_perc_by_month, short_perc_by_month]

  
    FG_perc_by_month = pd.DataFrame.from_dict(df_by_month, columns=["3-pointer", "mid-range", "short"], orient="index")
#     FG_perc_by_month.set_index("Month")
    FG_perc_by_month.index.name="Month"
    FG_perc_by_month.columns.name="Shot_Type"
    
    return FG_perc_by_month

ComputeGoalType(df_shots_player)
```

Output:
```
Shot_Type  3-pointer  mid-range     short
Month                                    
1           0.408696   0.453333  0.650000
2           0.469136   0.346939  0.648649
11          0.433628   0.506329  0.647059
12          0.352381   0.478873  0.704918
```

<!-- grade_id: cell-06c7dc215450e8bd, points: 25, grade: true, solution: false -->
```python
df_shots_player = GetPlayerData("stephen curry")
df_shots_player = AddMonth(df_shots_player)
df_shots_player = AddShotType(df_shots_player)
FG_perc_by_month = ComputeGoalType(df_shots_player)
assert [round(x, 2) for x in FG_perc_by_month["3-pointer"].tolist()] == [0.41, 0.47, 0.43, 0.35]
assert [round(x, 2) for x in FG_perc_by_month["short"].tolist()] == [0.65, 0.65, 0.65, 0.70]
```